# Problem 3

This problem has a single exercise worth a total of ten (10) points.

**Exercise 0** (10 points). Define a function, `UniqueCharacters(s)`, that given a string `s` returns a tuple of two elements: the first element is the number of unique alphabetic characters in the string and the second is the number of unique digits (base-10) in the string.

For example, the string `'ewwwffioj122434'` should output the following tuple: `(6, 4)`. The `6` occurs because there are 6 unique letters (`'e'`, `'w'`, `'f'`, `'i'`, `'o'`, and `'j'`) and `4` because there are 4 unique decimal digits (`'1'`,  `'2'`, `'3'`, `'4'`). Special characters may appear in the string but do not count as a letter or number.

```python
def UniqueCharacters(s):
    
    mySet = set(s)
    
    print(mySet)
    
    charCount = 0
    numCount = 0
    for element in mySet:
        print (element)
        if element.isalpha():
            charCount=charCount+1
        else:
            if element.isnumeric():
                numCount=numCount+1
            
    return charCount,numCount

```

<!-- grade_id: UniqueCharacters, points: 10, grade: true, solution: false -->
```python
# Test Cell: 'UniqueCharacters' (10 points)
assert UniqueCharacters('abc123') == (3,3)
assert UniqueCharacters('abc') == (3,0)
assert UniqueCharacters('aaa') == (1,0)
assert UniqueCharacters('aaa111') == (1,1)
assert UniqueCharacters('111') == (0,1)
assert UniqueCharacters('123') == (0,3)
assert UniqueCharacters('') == (0,0)
assert UniqueCharacters('///;;;...,????!!!!!###$$$%%%') == (0,0)
assert UniqueCharacters('//23bd') == (2,2)
assert UniqueCharacters('b2b3n4s4d9') == (4,4)
assert UniqueCharacters('9090909p0y90p90y90') == (2,2)
assert UniqueCharacters('ieowjfiojfioj2342io4ji') == (6,3)
assert UniqueCharacters('ewwwffioj122434') == (6,4)
assert UniqueCharacters('dwdj2ru3jf894jgf.,/23,4./3ei8fj2389ej89/,.,./2dd32je98dej89ij') == (9,5)
print("\n(Passed!)")
```

Output:
```
{'a', 'c', '3', '2', '1', 'b'}
a
c
3
2
1
b
{'a', 'b', 'c'}
a
b
c
{'a'}
a
{'a', '1'}
a
1
{'1'}
1
{'3', '2', '1'}
3
2
1
set()
{'!', '$', ';', '/', '%', '?', '.', '#', ','}
!
$
;
/
%
?
.
#
,
{'3', '/', '2', 'b', 'd'}
3
/
2
b
d
{'4', '3', '2', 's', 'n', 'b', '9', 'd'}
4
3
2
s
n
b
9
d
{'9', '0', 'p', 'y'}
9
0
p
y
{'4', 'e', '3', 'f', 'w', 'j', '2', 'i', 'o'}
4
e
3
f
w
j
2
i
o
{'4', 'e', '3', 'f', 'j', '2', 'i', '1', 'w', 'o'}
4
e
3
f
j
2
i
1
w
o
{'4', 'e', '8', '3', 'g', 'u', 'f', 'j', '/', '2', 'i', '.', 'w', 'r', '9', 'd', ','}
4
e
8
3
g
u
f
j
/
2
i
.
w
r
9
d
,

(Passed!)
```

**Fin!** You've reached the end of this problem. Don't forget to restart the
kernel and run the entire notebook from top-to-bottom to make sure you did
everything correctly. If that is working, try submitting this problem. (Recall
that you *must* submit and pass the autograder to get credit for your work!)